---
title       : Insert the chapter title here

--- type:NormalExercise lang:python xp: skills: key:2870874031
## <<<TITLE>>> 


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```
description : Insert the chapter description here